<a href="https://colab.research.google.com/github/2211Suraj/CE880_demo/blob/main/week7/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Suraj Rajan"
COLLABORATORS = "2111826"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling. Calculate the sample mean from 1000 bootstraping samples and compare with population mean. 

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=1000)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300

def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    sample_m = []
    x = sample_m.append(pop)
    y = np.mean(sample_m)
    return y
    raise NotImplementedError()
```

Output:
```
253.241
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

---
## Question 2: 

#### Majority Voting
Given a list = `[100,200,300,100,200,100,100,10,30,20,20,10]`. Write a function to find the majority voting.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
a = np.array([100,200,300,100,200,100,100,10,30,20,20,10])

def my_majorityvote():
    """Write a function to find the majority voting."""
    alfa = np.bincount(a)
    beta = np.max(alfa)
    for i in range(len(alfa)):
      if alfa[i] ==beta:
        return i

    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 100, rel_tol = 0.05)
```

---
## Question 3: 

#### Decision Tree
USe the `IRIS` dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely `IRIS Setosa`, `IRIS Versicolour`, and `IRIS Virginica`, with the following attributes: `sepal length`, `sepal width`, `petal length`, and `petal width`. Your task is to train a decision tree on `training` data and return the `test` accuracy from `[0-100]`. Hint: multiply for `accuracy_score` it with 100. 

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    ab = DecisionTreeClassifier()
    ab = ab.fit(X_train,y_train)
    y_pred = ab.predict(X.test)
    c = accuracy_score(y_test,y_pred)
    return c*100
    
    raise NotImplementedError()
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 94.736, rel_tol = 0.05)
```

Output:
```
<ipython-input-14-e0032349cbba> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_DTree(), 94.736, rel_tol = 0.05)

<ipython-input-11-fba52d0166bd> in my_DTree()
     24     ab = DecisionTreeClassifier()
     25     ab = ab.fit(X_train,y_train)
---> 26     y_pred = ab.predict(X.test)
     27     c = accuracy_score(y_test,y_pred)
     28     return c*100

AttributeError: 'numpy.ndarray' object has no attribute 'test'
```

```python

```