<a href="https://colab.research.google.com/github/2211Suraj/CE880_demo/blob/main/week10/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Suraj Rajan"
COLLABORATORS = "211826"
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 



```python
import numpy as np
from sklearn import datasets, linear_model
from sklearn.metrics import mean_squared_error, r2_score

# Load the diabetes dataset
diabetes_X, diabetes_y = datasets.load_diabetes(return_X_y=True)

# Use only one feature
diabetes_X = diabetes_X[:, np.newaxis, 2]

# Split the data into training/testing sets
diabetes_X_train = diabetes_X[:-20]
diabetes_X_test = diabetes_X[-20:]

# Split the targets into training/testing sets
diabetes_y_train = diabetes_y[:-20]
diabetes_y_test = diabetes_y[-20:]
```

---
## Question 1: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the `Ridge` linear regression model to get the test accuracy and return `mean squared error`.  

<!-- grade_id: my_reg_model_mse_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `mean_squared_error` on diabetes_X_test. 
    Please use Ridge() model to train a regressor. """
    
    log = linear_model.redge()
    log.fit(diabetes_X_train,diabetes_y_train)
    pred = log.predict(diabetes_X_train)
    MSE = mean_squared_error(diabetes_y_test,pred)
    return MSE
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_mse_test, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
                    3604.26, rel_tol = 0.05)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-8e05269cb419> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
      4                     3604.26, rel_tol = 0.05)

NameError: name 'diabetes_X_train' is not defined
```

---
## Question 2: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the `Ridge` linear regression model to get the test accuracy and return `r2` (r sqaured value).  

<!-- grade_id: my_reg_model_r2_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `r2_score` on diabetes_X_test data . 
    Please use LinearRegression model to train a regressor. """
    Reg = linear_model.Ridge()
    Reg.fit(daibetes_X_train,diaabetes_y_train)
    pred = Reg.predict(diabetes_x_test)
    score = r2_score(diabetes_y_test,pred)
    return score
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_r2_test, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
                                    diabetes_y_test), 0.2539, rel_tol = 0.05)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-4-ef50329c6566> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
      4                                     diabetes_y_test), 0.2539, rel_tol = 0.05)

NameError: name 'diabetes_X_train' is not defined
```